# COGS 108 - Assignment 1: Setting Up

### In markdown cell below, write in your Github username:

<!-- grade_id: git-user-ID, points: 0, grade: true, solution: true -->
My Github username is: 

Follow along on the instructions in the assignment description to do Part-1 on Github. 

<!-- grade_id: add2, points: 0, grade: false, solution: true -->
```python
# Write a function that adds 2 numbers. Call it 'add2'
### BEGIN SOLUTION
def add2(a, b):
    return a + b
### END SOLUTION
```

<!-- grade_id: add2-tests, points: 0, grade: true, solution: false -->
```python
assert add2(1, 2) == 3
assert add2(-1, 1) == 0
```