# COGS 108 - Assignment 1: Setting Up

### In markdown cell below, write in your Github username:

My Github username is: 

Follow along on the instructions in the assignment description to do Part-1 on Github. 

<!-- grade_id: add2, points: 0, grade: false, solution: true -->
```python
# Write a function that adds 2 numbers. Call it 'add2'
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: add2-tests, points: 1, grade: true, solution: false -->
```python
assert add2(1, 2) == 3
assert add2(-1, 1) == 0
```

<!-- grade_id: imports, points: 0, grade: false, solution: true -->
```python
# Import the five core python data-science modules, with their typical abbreviations
#  {numpy, scipy, pandas, matplotlib.pyplot, sklearn} as {np, sp, pd, plt, skl}
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: autograder-tests, points: 1, grade: true, solution: false -->
```python
assert(np)
assert(sp)
assert(pd)
assert(plt)
assert(skl)
```