# COGS 108 - Assignment 2: Data Exploration

# Important

- Rename this file to 'A2_A########.ipynb' (filled in with your student ID) before you submit it. Submit it to TritonED.
- This assignment has hidden tests: tests that are not visible here, but that will be run on your submitted. 
    - This means passing all the tests you can see in the notebook here does not guarantee you have the right answer!

## Part 1 - Data Wrangling

You are provided with two data files:
- age_steps.csv
- name_income_id.json

```python
# Imports
#  Note: these are all the imports you need! Do not import any other functions / packages
%matplotlib inline

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

<!-- grade_id: load_steps, points: 0, grade: false, solution: true -->
```python
# 1a) Load the csv file into a pandas dataframe. Call it 'df_steps'.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_load_steps, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(df_steps, pd.DataFrame)

```

<!-- grade_id: load_income, points: 0, grade: false, solution: true -->
```python
# 1b) Load the json file into a (different) pandas dataframe. Call it 'df_income'.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_load_income, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(df_income, pd.DataFrame)

```

<!-- grade_id: drop_extra_cols, points: 0, grade: false, solution: true -->
```python
# 1c) Drop the 'first_name' and 'last_name' columns from the df_income data frame.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_drop_ext_cols, points: 0.25, grade: true, solution: false -->
```python
assert 'first_name' not in df_income.columns
assert 'last_name' not in df_income.columns
```

<!-- grade_id: merge_dat, points: 0, grade: false, solution: true -->
```python
# 1d) Combine the two dataframes into one combined, new dataframe. Call it 'df'.
#  Note: For this part, it doesn't in which order you merge the dataframes.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_merge_dat, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(df, pd.DataFrame)
assert set(df.columns) == set(['id', 'income', 'age', 'steps'])
```

<!-- grade_id: reorder_cols, points: 0, grade: false, solution: true -->
```python
# 1e) Re-Order the Column ID's to be [id, age, steps, income]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_reorder_cols, points: 0.25, grade: true, solution: false -->
```python
assert list(df.columns) == ['id', 'age', 'steps', 'income']
```

## Part 2 - Data Cleaning

<!-- grade_id: check_n_nan, points: 0, grade: false, solution: true -->
```python
# 2a) First check how many missing data (NaNs) there are in the income data.
#  Save this number into a variable called 'n_nan'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_check_n_nan, points: 0.25, grade: true, solution: false -->
```python
assert(n_nan)

```

<!-- grade_id: rm_nans, points: 0, grade: false, solution: true -->
```python
# 2b) Remove all rows from df that have missing data (that have NaN values)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_nans, points: 0.5, grade: true, solution: false -->
```python
assert(sum(np.isnan(df['income'])) == 0)
```

## Data Visualization

- You will be making histograms, check out this example on how to do that.
    - http://matplotlib.org/1.2.1/examples/pylab_examples/histogram_demo.html
    - This example has more than you need: 
        - Titles, axis labels, etc are optional (but recommended)
        - You do no need to add a 'best fit' line.
- Everywhere that we ask you to create a plot, make sure to leave the 'plt.gcf()' line at the end of the cell. 
    - This is used to get a figure handle in order to programmatically evaluate your plots.
- For all your histograms, use 25 bins

<!-- grade_id: hist_age, points: 0, grade: false, solution: true -->
```python
# 2c) Plot the histogram of the age variable
# YOUR CODE HERE
raise NotImplementedError()
f1 = plt.gcf()
```

<!-- grade_id: test_hist_age, points: 0.25, grade: true, solution: false -->
```python
assert f1.gca().has_data()
```

<!-- grade_id: hist_steps, points: 0, grade: false, solution: true -->
```python
# 2d) Plot the histogram of the steps variable
# YOUR CODE HERE
raise NotImplementedError()
f2 = plt.gcf()
```

<!-- grade_id: test_hist_steps, points: 0.25, grade: true, solution: false -->
```python
assert f2.gca().has_data()
```

### M1 - What do you notice about steps data from the plot?

<!-- grade_id: notice_steps, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: check_n_neg, points: 0, grade: false, solution: true -->
```python
# 2e) Check how many -1 values there are in steps. Save the number into a variable called 'n_neg'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_check_n_neg, points: 0.25, grade: true, solution: false -->
```python
assert(n_neg)

```

<!-- grade_id: rm_neg, points: 0, grade: false, solution: true -->
```python
# 2f) Remove all rows where steps is equal to -1
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_neg, points: 0.5, grade: true, solution: false -->
```python
assert -1 not in df['steps']
```

<!-- grade_id: hist_income, points: 0, grade: false, solution: true -->
```python
# 2g) Plot the histogram of the income variable
# YOUR CODE HERE
raise NotImplementedError()
f3 = plt.gcf()
```

<!-- grade_id: test_hist_income, points: 0.25, grade: true, solution: false -->
```python
assert f3.gca().has_data()
```

### M2 - What do you notice about the income data from the plot?

<!-- grade_id: notice_income, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: recoded_zero, points: 0, grade: false, solution: true -->
```python
# 2h) For all rows in which income is equal to zero, recode income to be equal to 1
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_recoded_zero, points: 0.5, grade: true, solution: false -->
```python
assert sum(df['income'] == 0.) == 0
```

<!-- grade_id: log, points: 0, grade: false, solution: true -->
```python
# 2i) Transform income using a log10 transform
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_log, points: 0.5, grade: true, solution: false -->
```python
assert np.isclose(min(df['income']), 0.0, 0.1)
assert np.isclose(max(df['income']), 5.69, 0.1)
```

<!-- grade_id: replot_log, points: 0, grade: false, solution: true -->
```python
# 2j) Replot the histogram for income data after the data transformation
# YOUR CODE HERE
raise NotImplementedError()
f4 = plt.gcf()
```

<!-- grade_id: test_replot_hist, points: 0.25, grade: true, solution: false -->
```python
assert f4.gca().has_data()
```

### M3 - What do you notice about the income data after you replot it?

<!-- grade_id: notice_log2, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: rm_young, points: 0, grade: false, solution: true -->
```python
# 2k) Remove all rows where age is less than 18
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_young, points: 0.5, grade: true, solution: false -->
```python
assert min(df['age']) >= 18
```

### Removing outliers
Next, we're going to remove all rows where the data for the steps column exceeds +/- 3 standard deviations above or below the mean.

<!-- grade_id: rm_outliers_setup, points: 0, grade: false, solution: true -->
```python
# 2l) For the steps data: 
# - Calculate what three times the standard deviation is, and save it to variables called 'std3'. 
# - Using 'std3', create new variables, called std_max' and 'std_min'.
#    These will store our cut-off values: 3 std above and below the mean respectively. 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_rm_outliers_setup, points: 0.5, grade: true, solution: false -->
```python
assert(std3)
assert(std_min)
assert(std_max)

```

<!-- grade_id: rm_outliers, points: 0, grade: false, solution: true -->
```python
# 2m) Now use these values to remove any rows with steps data above or below these thresholds

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: tests_rm_outliers, points: 0.5, grade: true, solution: false -->
```python
assert min(df['steps']) > 3446
assert max(df['steps']) < 30594
```

```python
# Check that the following test passes - if it fails, you have the wrong amount of data left. 
#  If so, go back and try to fix your cleaning steps, else your analysis answers will be wrong.
assert df.shape == (879, 4)
```

## Part 3 - Basic Analyses

<!-- grade_id: calc_corrs, points: 0, grade: false, solution: true -->
```python
# 3a) Calculate the pairwise correlations between all variables. 
#  Notes: do this with a pandas function. Keep all columns (including ID).

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_calc_corrs, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(corrs, pd.DataFrame)
assert corrs.shape == (4, 4)
```

<!-- grade_id: plot_scatter, points: 0, grade: false, solution: true -->
```python
# 3b) Plot the data using pandas scatter_matrix, saving it to output variable name 'fd'.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_plot_scatter, points: 0.5, grade: true, solution: false -->
```python
fd[0][0].axes.has_data()
```

# Part 4 - Predictions

Notes:
- Use the polyfit function from np, in the way that we did in 'SectionMaterials/02-DataAnalysis
- The predictions should be in dollars (so you'll have to 'un-log' the prediction)

<!-- grade_id: model_1, points: 0, grade: false, solution: true -->
```python
# 4a) Use polyfit to fit a 1-degree linear model, predicting income from age.
#  Call the output parameters 'a1' and 'b1'
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_model_1, points: 0.5, grade: true, solution: false -->
```python
assert(a1)
assert(b1)

```

<!-- grade_id: pred_1, points: 0, grade: false, solution: true -->
```python
# 4b) Use the model parameters from 4a to predict the income of a 75 year old. Call your prediction 'pred1'.
#  Note: Convert the prediction to be in real dollars (not log dollars)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_pred_1, points: 0.5, grade: true, solution: false -->
```python
assert(pred1)

```

<!-- grade_id: model_2, points: 0, grade: false, solution: true -->
```python
# 4c) Use polyfit to fit a 1-degree linear model, predicting income from steps. 
#  Call the output parameters 'a2' and 'b2'
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_model_2, points: 0.25, grade: true, solution: false -->
```python
assert(a2)
assert(b2)

```

<!-- grade_id: pred_2, points: 0, grade: false, solution: true -->
```python
# 4d) Predict the income of someone who took 20,000 steps. Call your prediction 'pred2'.
#  Note: Convert the prediction to be in real dollars (not log dollars)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_pred_2, points: 0.25, grade: true, solution: false -->
```python
assert(pred2)

```

<!-- grade_id: better_model, points: 0, grade: false, solution: true -->
```python
# 4e) Based on your analysis in Part 3, which model do you think is more likely to have made a good prediction?
#  Create a variable called 'better_model', and set it equal to 'age' or 'steps', whichever your answer is.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_better_model, points: 0.5, grade: true, solution: false -->
```python
assert better_model
```

### M4 - Why do you think (whichever model you answered) is better?

<!-- grade_id: model_ans, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Done! Upload this notebook to TritonED