# COGS 108 - Assignment 1: Setting Up

## IMPORTANT
### Rename this file to 'A1_A########.ipynb' (filled in with your student ID before you submit it)

### In the markdown cell below, write in your Student ID and Github username. 

My Student ID number is: 

My Github username is: 

## Part 1

Follow along on the instructions in the assignment description to do Part-1 on Github. 

```python
# Run this cell to check your version of IPython. If the message prints out, stop and upgrade Jupyter.
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Part 2

The Python portion of this assignment is just to test that the you have everything set up and that the submission and grading system all works. 

Whenever you see:

YOUR CODE HERE <br>
raise NotImplementedError()

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise this error and fail the tests.

Note that in this notebook all the tests are public - that is, you can see that tests that will be used to evaluate your code, and so can check your answers. This will not always be the case - future assignments will also include hidden tests. 

<!-- grade_id: add2, points: 0, grade: false, solution: true -->
```python
# Write a function that takes two inputs, adds them together, and returns the result. Call it 'add2'
### BEGIN SOLUTION
def add2(a, b):
    return a + b
### END SOLUTION
```

<!-- grade_id: add2-tests, points: 2, grade: true, solution: false -->
```python
# These are the test for the question above
assert add2(1, 2) == 3
assert add2(-1, 1) == 0
```

<!-- grade_id: imports, points: 0, grade: false, solution: true -->
```python
# Import the five core python data-science modules, with their typical abbreviations
#  {numpy, scipy, pandas, matplotlib.pyplot, sklearn} as {np, sp, pd, plt, skl}
### BEGIN SOLUTION
import numpy as np
import scipy as sp
import pandas as pd
import matplotlib.pyplot as plt
import sklearn as skl
### END SOLUTION
```

<!-- grade_id: autograder-tests, points: 2, grade: true, solution: false -->
```python
# These are the test for the question above
assert(np)
assert(sp)
assert(pd)
assert(plt)
assert(skl)
```